## Overview of Project 2: Lettuce with Figs

Package name: `edu.colorado.csci3155.project2`
Main class: `edu.colorado.csci3155.project2.CanvasRenderTest`



From the instructions in the project folder (`/doc/LettuceWithFigs-Project2.ipynb`):

###  Task 1
       Implement the canvas operations using the data structures defined in the file MyCanvas.scala. Unit tests provided in CanvasTests.scala.
       
###  Task 2
       Fill in the missing code for the interpreter (Interpreter.scala). The value types are in Value.scala, and the environment is defined in Environment.scala.

    You can throw an IllegalArgumentException whenever you encounter an error. A message describing the error should be included in the exception.
       
      

## Continuation Passing Style

Continuation passing style (CPS) makes use of first-class functions by requiring that all functions take an extra argument `k`, which we call the continuation. `k` must be a function, and when a function returns, it calls `k` on the return value and returns that instead.

```scala
// This function takes an integer x and returns x + 1
def addOne(x: Int): Int = {
    val result = x + 1
    result
}
```

Output:
```
defined [32mfunction[39m [36maddOne[39m
```

```scala
addOne(5)
```

Output:
```
[36mres1[39m: [32mInt[39m = [32m6[39m
```

```scala
// In CPS, we would write...
def addOne_cps(x: Int, k: Int => Int): Int = {
    val result = x + 1
    k(result)
}
```

Output:
```
defined [32mfunction[39m [36maddOne_cps[39m
```

```scala
val identity: Int => Int = y => y
addOne_cps(5, identity)
```

Output:
```
[36midentity[39m: [32mInt[39m => [32mInt[39m = ammonite.$sess.cmd3$Helper$$Lambda$1851/0x0000000801ac7840@26138d99
[36mres3_1[39m: [32mInt[39m = [32m6[39m
```

```scala
// Use generics to allow any return type for k
def addTwo_cps[T](x: Int, k: Int => T): T = {
    val result = x + 2
    k(result)
}
```

Output:
```
defined [32mfunction[39m [36maddTwo_cps[39m
```

```scala
addTwo_cps(5, identity)
```

Output:
```
[36mres5[39m: [32mInt[39m = [32m7[39m
```

```scala
addTwo_cps(5, x => List(x))
```

Output:
```
[36mres6[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m7[39m)
```

## Exercise: Side Effects
Given the following functions, change them to use continuations:

```scala
// Function with side effects
def printFive(f: Int => String): Int = {
    val five = 5
    println(f(five))
    five
}
```

Output:
```
defined [32mfunction[39m [36mprintFive[39m
```

```scala
def printFive_cps(f: Int => String, k: Int => Int): Int = {
    // Your Code
    val result = 5
    println(f(result))
    k(result)
}
```

Output:
```
defined [32mfunction[39m [36mprintFive_cps[39m
```

```scala
printFive(x => x.toString)
```

Output:
```
5
[36mres12[39m: [32mInt[39m = [32m5[39m
```

```scala
printFive_cps(x => x.toString, y => y * 10)
```

Output:
```
5
[36mres13[39m: [32mInt[39m = [32m50[39m
```

### Exercise: Backtracking
Search a binary tree using CPS. Return true if the tree has a node with the integer `x` as a value.

<!-- grade_id: cell-bdaf875cad5b0678, points: 0, grade: false, solution: true -->
```scala
sealed trait Tree
case object Empty extends Tree
case class Node(left: Tree, value: Int, right: Tree) extends Tree

def search(t: Tree, x: Int, fail_continuation: () => Boolean): Boolean =
    // Your Code
    t match {
        case Empty => fail_continuation()
        case Node(l, v, r) if (v == x) => true
        case Node(l, v, r) => search(l, x, () => search(r, x, fail_continuation))
    }

```

Output:
```
defined [32mtrait[39m [36mTree[39m
defined [32mobject[39m [36mEmpty[39m
defined [32mclass[39m [36mNode[39m
defined [32mfunction[39m [36msearch[39m
```

<!-- grade_id: cell-df2188387bae0d26, points: 1, grade: true, solution: false -->
```scala
val t = Node(Empty, 10, Node(Empty, 6, Empty))
assert(search(t, 10, () => false))
assert(!search(t, 0, () => false))
```

Output:
```
[36mt[39m: [32mNode[39m = [33mNode[39m(Empty, [32m10[39m, [33mNode[39m(Empty, [32m6[39m, Empty))
```

```scala

```